# Working with [ICD9](https://www.cdc.gov/nchs/icd/icd9.htm) Data

```python
%matplotlib inline
```

```python
import pickle
from collections import defaultdict
from nose.tools import assert_equals
```

```python
with open("icd9.pickle","rb") as f0:
    icd9_data = pickle.load(f0)
```

## What does our data look like?

```python
icd9_data[:5]
```

Output:
```
((56, '198.3', 'SECONDARY MALIGNANT NEOPLASM OF BRAIN AND SPINAL C'),
 (56, '162.8', 'MALIGNANT NEOPLASM OF OTHER PARTS OF BRONCHUS OR L'),
 (56, '531.40', 'CHRONIC OR UNSPECIFIED GASTRIC ULCER WITH HEMORRHA'),
 (56, '276.1', 'HYPOSMOLALITY AND/OR HYPONATREMIA'),
 (56, '428.0', 'CONGESTIVE HEART FAILURE UNSPECIFIED'))
```

```python
for d in icd9_data:
    try:
        if "HEP" in d[2]:
            print(d)
    except Exception as Error:
        pass
```

Output:
```
(78, '070.54', 'CHRONIC HEPATITIS C WITHOUT HEPATIC COMA')
(106, 'V02.62', 'CARRIER OR SUSPECTED CARRIER OF HEPATITIS C')
(124, '070.70', 'UNSPECIFIED VIRAL HEPATI                \r')
(117, '070.54', 'CHRONIC HEPATITIS C WITHOUT HEPATIC COMA')
(117, '070.44', 'CHR HEPAT C W/ HEP COMA                 \r')
(117, '572.4', 'HEPATORENAL SYNDROME                    \r')
(390, '070.70', 'UNSPECIFIED VIRAL HEPATI                \r')
(433, '070.32', 'CHRON VIR HEP B W/O COMA                \r')
(433, '070.54', 'CHR HEPAT C W/O HEP COMA                \r')
(433, '070.30', 'VIR HEP B W/O COM, HEP D                \r')
(433, '070.70', 'UNSPECIFIED VIRAL HEPATI                \r')
(590, '571.1', 'ACUTE ALCOHOLIC HEPATITIS')
(634, '572.4', 'HEPATORENAL SYNDROME                    \r')
(940, '070.54', 'CHRONIC HEPATITIS C WITHOUT HEPATIC COMA')
(940, '155.1', 'MALIGNANT NEOPLASM OF INTRAHEPATIC BILE DUCTS')
(1106, '070.32', 'CHRON VIR HEP B W/O COMA                \r')
(1100, '070.54', 'CHRONIC HEPAT
```

### Each element of `icd9_data` is a tuple with three elements
1. A patient id
1. An ICD9 code
1. The label for that ICD9 code

#### Create a dictionary named `icd9_map` with keys ICD9 code and values the ICD9 label

strip extra white spaces

<!-- grade_id: cell-5d44913c5eba6a2c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
icd9_map = {}
for d in icd9_data:
    try:
        key = d[1]
        val = " ".join(d[2].split())
        icd9_map[key] = val
        
    except Exception as Error:
        pass
   
print (icd9_map)
#raise NotImplementedError()
```

Output:
```
{'861.21': 'CONTUSION OF LUNG WITHOUT OPEN WOUND INTO THORAX', 'V10.00': 'HX OF GI MALIGNANCY NOS', '305.1': 'TOBACCO USE DISORDER', '196.1': 'SECONDARY AND UNSPECIFIED MALIGNANT NEOPLASM OF IN', '428.23': 'ACUTE ON CHR SYSTO HRT', '415.19': 'OTHER PULMONARY EMBOLISM AND INFARCTION', '611.0': 'INFLAMMATORY DISEASE OF BREAST', '568.81': 'HEMOPERITONEUM (NONTRAUMATIC)', '070.51': 'ACUTE OR UNSPECIFIED HEPATITIS C WITHOUT HEPATIC C', '518.5': 'PULMONARY INSUFFICIENCY FOLLOWING TRAUMA AND SURGE', '203.00': 'MULTIPLE MYELOMA WITHOUT REMISSION', 'V10.46': 'PERSONAL HISTORY OF MALIGNANT NEOPLASM OF PROSTATE', '438.89': 'OTHER LATE EFFECTS OF CEREBROVASCULAR DISEASE', '335.20': 'AMYOTROPHIC LATERAL SCLEROSIS', '423.8': 'OTHER SPECIFIED DISEASES OF PERICARDIUM', '560.2': 'VOLVULUS', '747.0': 'PATENT DUCTUS ARTERIOSUS', '441.1': 'THORACIC ANEURYSM RUPTURED', '053.9': 'HERPES ZOSTER NOS', '573.3': 'HEPATITIS UNSPECIFIED', '812.00': 'FRACTURE OF UNSPECIFIED PART OF UPPER END OF HUMER', '769': 'RES
```

```python
icd9_patients = defaultdict(list)
patient_diagnoses = defaultdict(list)
```

### How many patients are there for each diagnosis?

Create a list named `icd9_patients_list` sorted by the number of patients per diagnosis.

<!-- grade_id: cell-ae2a3dc827383e7f, points: 0, grade: false, solution: true -->
```python
icd9_patients = defaultdict(list)
# YOUR CODE HERE
icd9_patients_list = []
for d in icd9_data:
    if dicd9_patients_list


#raise NotImplementedError()
```

Output:
```
defaultdict(<class 'list'>, {})
```

```python
import random
d,p = random.choice(icd9_patients_list)
print(d,len(p),sep="\n")
print(icd9_patients["V12.59"])
for d,p in icd9_patients.items():
    if "HX" in d:
        print(d,len(p))
```

```python
assert_equals(len(icd9_patients["V12.59"]),5)
assert_equals(len(icd9_patients["572.2"]),12)
```

#### Loop through  `icd9_patients_list`
1. for each element in `icd9_patients_list` print the ICD9 label corresponding to the code and the number of patients with each diagnosis.

<!-- grade_id: cell-121dd04abe83faee, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### How many diagnoses does each patient have?

<!-- grade_id: cell-305cd7b053c0116a, points: 0, grade: false, solution: true -->
```python
patient_diagnoses = defaultdict(list)
# YOUR CODE HERE
raise NotImplementedError()
```

```python
for p, d in patient_diagnoses_list:
    print(p)
    print(d) # replace with len(d)
    print("\n")
```

<!-- grade_id: cell-322e571e6bd058c0, points: 0, grade: true, solution: false -->
```python
assert_equals(len(patient_diagnoses[2512]),49)
assert_equals(len(patient_diagnoses[353]),56)
assert_equals(len(patient_diagnoses_list[0][1]),125)
assert_equals(len(patient_diagnoses_list[45][1]),23)

```

```python

```